## Exercise 03.1

Compare the computed values of 

$$ 
d_0 = a \cdot b + a  \cdot  c
$$

and 

$$ 
d_1 = a \cdot (b + c)
$$

when $a = 100$, $b = 0.1$ and $c = 0.2$. Store $d_{0}$ in the variable `d0` and $d_{1}$ in the variable `d1`.


Try checking for equality, e.g. `print(d0 == d1)`. 

<!-- grade_id: cell-447110c225124fc4, points: 0, grade: false, solution: true -->
```python
d0=100*0.1+100*0.2
d1 = 100*(0.1+0.2)
print(d0==d1)
print(d0)
```

Output:
```
False
30.0
```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

## Exercise 03.2

For the polynomial 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
compute $f$ using: (i) the compact form $(x + y)^{6}$; and (ii) the expanded form for:

(a) $x = 10$ and $y = 10.1$

(b) $x = 10$ and $y = -10.1$

and compare the number of significant digits for which the answers are the same.
Store the answer for the compact version using the variable `f0`, and using the variable `f1` for the expanded version.

For case (b), compare the computed and analytical solutions and consider the relative error.
Which approach would you recommend for computing this expression?

#### (a) $x = 10$ and $y = 10.1$

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = 10.1
f0 = (x+y)**6

f1 = x**6+6*(x**5)*(y)+15*(x**4)*(y**2)+20*(x**3)*(y**3)+15*(x**2)*(y**4)+6*x*(y**5)+y**6
print(f0, f1)
#same value to 14 significant figures
```

Output:
```
65944160.60120103 65944160.601201
```

<!-- grade_id: cell-55d184514edf3a67, points: 0, grade: true, solution: false -->
```python
import math
assert math.isclose(f0, 65944160.60120103, rel_tol=1e-10)
assert math.isclose(f1, 65944160.601201, rel_tol=1e-10)
```

#### (b) $x = 10$ and $y = -10.1$

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = -10.1

f0 = (x+y)**6

f1 = x**6+6*(x**5)*(y)+15*(x**4)*(y**2)+20*(x**3)*(y**3)+15*(x**2)*(y**4)+6*x*(y**5)+y**6
print(f0, f1)
#actual value should be (-0.1)*6 = 1x 10^-6
#Therefore, f0 is closer-> the computed solution is more accurate here.
```

Output:
```
9.999999999999788e-07 9.958166629076004e-07
```

```python
import math
assert math.isclose(f0, 1.0e-6, rel_tol=1e-10)
assert math.isclose(f1, 1.0e-6, rel_tol=1e-2)
```

## Exercise 03.3

Consider the expression

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

When $x$ is very large, the denominator approaches zero, which can cause problems.

Try rephrasing the problem and eliminating the fraction by multiplying the numerator and denominator by $\sqrt{x^2 - 1} + x$ and evaluate the two versions of the expression when:

(a) $x = 1 \times 10^{7}$

(b) $x = 1 \times 10^{9}$ (You may get a Python error for this case. Why?)

#### (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
#import numpy as np
#print(np.finfo(f))
#64bits used to store floats here, just checking

x = 1.0e7
f = 1/((x**2-1)**0.5-x)

f1 = ((x**2-1)**0.5+x)/(-1)


denom = (x**2-1)**0.5-x #notzero
print(denom)
print(f)
print(f1)

```

Output:
```
-5.029141902923584e-08
-19884107.85185185
-19999999.999999948
```

#### (b) $x = 1 \times 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
import decimal

x = 10**9

f1 = ((x**2-1)**0.5+x)/(-1)
print(f1)




f = 1/((x**2-1)**0.5-x)

#f, the original form of the expression, gives a Zero division error because the denominator is so close to zero. 
#This happened because two values that are almost identical are subtracted from one another.
#whereas the adjusted form f1 avoids this; by rationalising the denominator, the numerator then becomes a sum of two positive values, rendering a zero divison error impossible.
```

Output:
```
-2000000000.0
ZeroDivisionError                         Traceback (most recent call last)
<ipython-input-1-ac63f0d422ab> in <module>
      9 
     10 
---> 11 f = 1/((x**2-1)**0.5-x)
     12 
     13 #f, the original form of the expression, gives a Zero division error because the denominator is so close to zero. Because two values that are almost identical are subtracted from one another.

ZeroDivisionError: float division by zero
```

```python
#exploring cause of the zero division error
#first we look at the denominator of f
denom = (x**2-1)**0.5-x
print(type(denom))
#denominator of f should actually be -5 x 10^-10, but here it is calculated to be zero.
#btw it's a float

print("this is the denominator as calculated for the original f")
print('{0:.30f}'.format(x**2-1))
#we print this value to suppsoedly 30 significant places and find that it has been rounded.
#This is what causes equivalency between the two terms in the denominator, and therefore the zero divsion error in f.


# on a side note, 0.5 does not have roundoff error here.


xsquared = decimal.Decimal(x**2)
xsquaredminone = xsquared -1
print("this is the denominator once decimal.Decimal is applied")
print(xsquaredminone)
pointfive = decimal.Decimal(0.5)
rooted = xsquaredminone **pointfive
xdec = decimal.Decimal(x)
result = 1/(rooted-xdec)
print("This is the final result:",result)

#Here I replicate the exact operations in f step by step but apply decimal.Decimal as above to prevent premature rounding 
#This allows f to give the same result as f1

#this tells us the key error with f is that x**2-1 is rounded to x**2. This results in a denomiantor of 0 and therefore zero error.
#the default level of precision for floats in depends on what hardware you're running

print((1e9)**2-1==1000000000000000000)
print((1e9)**2-1==999999999999999999)
#this is a much simpler demonstration of the core problem. 
print((1e7)**2-1)
print((1e8)**2-1)
#this issue did not occur earlier with 1e7 which had 14 significant figures
#but with 1e8, which gives 16 significant figures, the precision issue begins.
#Python uses 64 bit floats
#IEEE-754 standard means that bits are used to store the sign, exponent and then the fraction-> fraction is essentially the numeb rof significant figures.
#As such, Python actually has 53 bit precision for the fraction part.
#2^53 = 10^15.955
#Therefore precision cannot exceed 15.955 decimal digits.
#double precision floating point format

```

Output:
```
<class 'float'>
this is the denominator as calculated for the original f
1000000000000000000.000000000000000000000000000000
this is the denominator once decimal.Decimal is applied
999999999999999999
This is the final result: -2E+9
True
False
99999999999999.0
1e+16
```

```python
#to get greek letters, do \lambda then press tab key.
```