## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
#coordlist is a list of coordinate tuples

def area(coordlist):
    for i in coordlist:
        if len(i) !=2:
            raise SystemExit("invalid 2d coordinates")
            

        elif (isinstance(i[0], int) == False and isinstance(i[0], float) == False) or (isinstance(i[1],int) == False and isinstance(i[1], float)==False):
            raise SystemExit("coordinates not numbers")
            

    return abs(0.5*(coordlist[0][0]*(coordlist[1][1]-coordlist[2][1])+coordlist[1][0]*(coordlist[2][1]-coordlist[0][1])+coordlist[2][0]*(coordlist[0][1]-coordlist[1][1])))
            

coordlist = [(1,3),(4,5),(6,7)] #user could input tuples or just a list of tuples for the coordinates
#but here I've opted for a list of tuples as to minimise the arguments passed to the function for a simpler interface

#The for loop goes through each set of coordinates, checking for errors.
#the system exit error prevents the area calculation being repeated from within the for loop and raising more errors by breaking the loop.

print("The area is", (area(coordlist)))



```

Output:
```
The area is 1.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python

def area(vertices): #vertices is a list of coordinate tuples
    areasum = 0 # variables within functions only have local scope so must define within here
    vertices.append(vertices[0])
    for i in range(len(vertices)-1): #we need to repeat the final xn, yn but we don't then want to sum that one.
        areasum+= vertices[i][0]*vertices[i+1][1]-vertices[i+1][0]*vertices[i][1]
    return abs(0.5*areasum)

vertices = [(0,0),(4,0),(4,4),(2,5),(0,4)]
print(area(vertices))
#bwahahaha it works
        
```

Output:
```
18.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y): #vectors in list form
    "Return sum of two vectors"
    if len(x)!=len(y):
        raise ValueError("vectors not the same length")
    else:
        new_vector = [0]*len(x)
        for i in range(len(x)):
            new_vector[i]= x[i]+y[i]
            
        return new_vector

print(sum_vector([3,4,6.5],[5,7,9]))
            
            
```

Output:
```
[8, 11, 15.5]
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
#using the vectors a and b as above
sum_vector = [x + y for x, y in zip(a,b)]
 
#The zip() function returns a zip object
#this is an iterator (something that can be iterated through) of tuples
#where the first item in each passed iterator is paired together, then the second items are paired together etc.

print(sum_vector)

lionel = [(x,y) for x,y in zip(a,b)] #each element in the zip object is of the format (x,y)
print(lionel) #this is what the zip object looks like when converted into a list
#or just use list function
print(list(zip(a,b)))
```

Output:
```
[-2, 11.3, -20, 8]
[(0, -2), (4.3, 7), (-5, -15), (7, 1)]
[(0, -2), (4.3, 7), (-5, -15), (7, 1)]
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
colleges = {"Gonville and Caius": "CAI", "Girton": "G", "Clare":"CL", "King's":"K","Trinity":"T"}

colleges_inverse = {}
for i,j in colleges.items(): #.items() returns the dictionary's key value pairs
    
    colleges_inverse[j] = i #now j is the key, and i is the value
    #no indexing in dictionaries- instead, the key acts as an index.
    #we can create new entries in dictionaries just by declaring like this.
    
print(colleges_inverse)
sorted_colleges = ([i for i in colleges_inverse]) #only taking the abbreviations. don't use.sort()!

sorted_colleges.sort()
print(sorted_colleges)

#could just use sorted(....) which returns a new list
#.sort() modifies list in placee
#i.e. .sort() replaces original list, whereas sorted() returns a sorted copy without changing the original list.

#hence, using .sort() when defining a variable returns none because the actual list hasn't been stored yet in the first place.
#e.g. sorted_colleges = ([i for i in colleges_inverse]).sort() DOESN'T WORK
```

Output:
```
{'CAI': 'Gonville and Caius', 'G': 'Girton', 'CL': 'Clare', 'K': "King's", 'T': 'Trinity'}
['CAI', 'CL', 'G', 'K', 'T']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
college_dict = {"Gonville and Caius":{"abbrv": "CAI", "date":1348, "students":849}, "Clare":{"abbrv":"CL", "date":1326, "students":808}, "Girton":{"abbrv":"G", "date":1869, "students":808}, "King's":{"abbrv":"K","date":1441, "students":726},"Trinity":{"abbrv":"T","date":1546, "students":1054}}
#a dictionary that maps college names (the key) to dictionaries of abbreviation, foundation year, number of students
#so effectively a dictionary of dictionaries for each college containing the attribute key and the value

max_students = 0
oldest_date = 100000 #oldest college has the lowest number
for i,j in college_dict.items(): #j is the nested dictionary
    if  j["students"]>max_students: #can't use indices with dictionaries, must use key to get value!
        max_students = j["students"]
        max_st_abbrv = j["abbrv"]
        
    if j["date"] <oldest_date:
        oldest_date = j["date"]
        oldest_students = j["students"]
        oldest_abbrv = j["abbrv"]
    

print("College with most students:",max_st_abbrv, max_students, "students")
print("Oldest College:",oldest_abbrv, oldest_students, "students,", oldest_date, "founding year")
print("")

#Alternate method: you can pass a key to the sorted() function! This allows you to sort by a certain key within the nested dictionary.
#The key function returns a tuple containing what you wish to sort on. Or use a lambda function.

#specify a function to be called on each list element prior to making comparisons. Like a proxy.
def keyfunc(tup):
    key, college_dict = tup 
    return -college_dict["students"] #minus sign to reverse order (number of students)

blab = sorted(college_dict.items(), key = keyfunc)
#.items() returns tuples that keyfunc then acts on.
#two variables used to store the tuple
#key would be Gonville & Caius, college_dict would be the corresponding nested dictionary.
print(blab)
print("")
#here, a lambda function is used instead

itemstest = sorted(college_dict.items(), key=lambda x: (x[1]['students']),reverse = True) 
#x is each element of the list/dict
#.items() will return key, value pairs
#so x[1] gives the nested dictionary for each college YEET
print(itemstest)

maxstudents = sorted(college_dict, key=lambda x: (college_dict[x]['students']),reverse = True) #x is what the lambda function takes as input
#reverse to display desired element as the first value of the sorted list.

#key function applied to every list/dict element
#^here, the list elements are Gonville and Caius, King's, Trinity, Clare, Girton etc....
#therefore, college_dict[x] returns the corresponding dictionary!

oldestcollege = sorted(college_dict, key=lambda x: (college_dict[x]['date']))
print(maxstudents)
print(oldestcollege)

print("Most students",maxstudents[0],college_dict[maxstudents[0]]["abbrv"],college_dict[maxstudents[0]]["students"])
print("Oldest college",oldestcollege[0], college_dict[oldestcollege[0]]["abbrv"], college_dict[oldestcollege[0]]["students"],college_dict[oldestcollege[0]]["date"])
```

Output:
```
College with most students: T 1054 students
Oldest College: CL 808 students, 1326 founding year

[('Trinity', {'abbrv': 'T', 'date': 1546, 'students': 1054}), ('Gonville and Caius', {'abbrv': 'CAI', 'date': 1348, 'students': 849}), ('Clare', {'abbrv': 'CL', 'date': 1326, 'students': 808}), ('Girton', {'abbrv': 'G', 'date': 1869, 'students': 808}), ("King's", {'abbrv': 'K', 'date': 1441, 'students': 726})]

[('Trinity', {'abbrv': 'T', 'date': 1546, 'students': 1054}), ('Gonville and Caius', {'abbrv': 'CAI', 'date': 1348, 'students': 849}), ('Clare', {'abbrv': 'CL', 'date': 1326, 'students': 808}), ('Girton', {'abbrv': 'G', 'date': 1869, 'students': 808}), ("King's", {'abbrv': 'K', 'date': 1441, 'students': 726})]
['Trinity', 'Gonville and Caius', 'Clare', 'Girton', "King's"]
['Clare', 'Gonville and Caius', "King's", 'Trinity', 'Girton']
Most students Trinity T 1054
Oldest college Clare CL 808 1326
```

```python

```