## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    if n<0:
        raise ValueError("Please input a positive integer") #should be RecursionError instead
    elif n ==0:
        return 0
    elif n ==1:
        return 1
    else:
        return f(n-1) + f(n-2)
f(-10)
```

Output:
```
<ipython-input-31-4a06c9a867e0> in <module>
      9     else:
     10         return f(n-1) + f(n-2)
---> 11 f(-10)

<ipython-input-31-4a06c9a867e0> in f(n)
      2     "Compute the nth Fibonacci number using recursion"
      3     if n<0:
----> 4         raise ValueError("Please input a positive integer") #should be RecursionError instead
      5     elif n ==0:
      6         return 0

ValueError: Please input a positive integer
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    "Evaluate polynomial function"
    return x**3 - 6*x**2 + 4*x + 12

def compute_root(f, x0, x1, tol, max_it,it_count = 0):
    
    
    if max_it<0 or tol<0:
        raise ValueError
    if isinstance(max_it,int) == False:
        
        raise TypeError("Please input max iterations as an integer")
    if isinstance(x0,(int,float))== False:
        raise TypeError("Starting point not a number")
    
    if isinstance(x1,(int,float))==False:
        raise TypeError("Endpoint not a number")
    if isinstance(tol,(int,float))==False:
        raise TypeError("Tolerance not a number")
    if isinstance(f(x0),(int,float))==False:
        raise TypeError("Error with function, or f(x0) is undefined")
#x0 does not necessarily have to be less than x1 for this to work
#could include f(x0)>0 and f(x1)<0 or vice versa-....but as with the case x0=x1, this would be caught by the max iterations error.
        
    
    x_mid = (x0+x1)/2
    f_mid = f(x_mid) #original f_mid
    f0 = f(x0)

    if f0*f_mid<0: 
        x1 = x_mid #set the new endpoint as x_mid; not the same as x_mid = x1 !!!!
    else:
        x0 = x_mid #set the new starting point as x_mid
    it_count +=1 

    f_mid = f(x_mid) #new f_mid
    
    if it_count > max_it: 
        raise RuntimeError("Maximum iterations exceeded.") #perhaps the answer diverges, or more iterations are needed.
    elif abs(f_mid) < tol:
        return x_mid, f_mid, it_count
        #return ends the function; the iteration error takes priority so that goes first.
    else:
        return compute_root(f, x0, x1, tol, max_it, it_count) #recursion with new x0, x1
    #and itcount is no longer its default value

#23 iterations max_it will work, but 22 won't. So the error handling system is effective.
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=20)
```

```python

```