## Exercise 12.1

Create a class to represent vectors of arbitrary length and which is initialised with a list of values, e.g.:
```python
x = MyVector([0, 2, 4])
```

Equip the class with methods that:

1. Return the length of the vector (use name `size`)
2. Compute the norm of the vector $\sqrt{x \cdot x}$ (use name `norm`)
3. Compute the dot product of the vector with another vector (use name `dot`)

Test your implementation using two vectors of length 3. To help you get started, a skeleton of the class is provided below. Don't forget to use `self` where necessary.

```python
class MyVector:
    """A vector object that can return its size and norm, and can compute the dot product 
    with another vector  """
    
    def __init__(self, x):
        self.x = x #as an attribute, x is the data stored in the object
        
    # Return length of vector
    def size(self):

    # This allows access by index, e.g. y[2]
    def __getitem__(self, index):
        return self.x[index]

    # Return norm of vector
    def norm(self):
        
    
    # Return dot product of vector with another vector
    def dot(self, other):
        
#proper code is below
```

### Solution

<!-- grade_id: cell-ff4d6cba19bb1b8e, points: 0, grade: false, solution: true -->
```python
import math  # Use this to get the sqrt

class MyVector:
    """A vector object that can return its size and norm, and can compute the dot product 
    with another vector  """

    def __init__(self, x):
        self.x = x #x is essentially the input
        
    # Return length of vector
    def size(self):
        s = 0
        for i in self.x:
            s+=1
        return s
    
    # This allows access by index, e.g. y[2]
    def __getitem__(self, index): #getitem is one of those magic methods.
        return self.x[index] #x is a list, so return the index using that.

    # Return norm of vector
    def norm(self):
        n = 0
        for i in self.x:
            n+=i**2
        return math.sqrt(n)
    
    # Return dot product of vector with another vector
    def dot(self, other):
        d = 0
        for i in range(self.size()):
            d+=(self[i]*other[i])
        return d
    
#self is the MyVector object
#self.x is the attribute data stored- which (in this example) is a list. Hence self.x does not have the attribute .size()
#but self does have that attribute.
```

```python
k = MyVector([3,4,5])
print(k.size())
print(k.norm())
w = MyVector([6,7,8])
print(k.dot(w))
#it works!
print(k[0])
```

Output:
```
3
7.0710678118654755
86
3
```

<!-- grade_id: cell-b601b5cc8cf8ab44, points: 0, grade: true, solution: false -->
```python
# Create two vectors
u = MyVector([1, 1, 2])
v = MyVector([2, 1, 1])

assert u.size() == 3
assert round(u.norm() - 2.449489742783178) == 0.0
assert round(u.dot(v) - 5.0, 10) == 0.0 
```

## Exercise 12.2

1. Create a class for holding a student record entry. It should have the following attributes:
   - Surname
   - Forename
   - Birth year
   - Tripos year
   - College
   - CRSid (optional field)
1. Equip your class with the method '`age`' that returns the age of the student in whole years
1. Equip your class with the method '`__repr__`' such using `print` on a student record displays with the format

       Surname: Bloggs, Forename: Andrea, College: Churchill

1. Equip your class with the method `__lt__(self, other)` so that a list of record entries can be sorted by 
   (surname, forename) using the Python built-in sort function. 
   
   Create a list of entries and test the sorting. Make sure you have two entries with the same
   surname.

> Recall that the methods starting with `__`, e.g. `__lt__` and `__repr__`, should **not** be called directly. 
> Python with map them to other operations, e.g. `__repr__` is called when using `print`, and `__lt__` is called 
> when using `<`. These functions must have a return value.

*Hint:* To get the current year:

```python
import datetime
year = datetime.date.today().year
print(year)
```

Output:
```
2020
```

### Solution

<!-- grade_id: cell-6e857f6089fe0ba8, points: 0, grade: false, solution: true -->
```python
import datetime 
class StudentEntry:
    def __init__(self, surname, forename, birth_year, year, college, crsid=None):
        self.surname = surname
        self.forename = forename #all attributes, can't be called, so no braces at the end
        self.birth_year = birth_year
        self.year = year
        self.college = college
        self.crsid = crsid

    def age(self):
        return datetime.date.today().year - self.birth_year
    
    def __repr__(self): #what's shown when you print the object
        
        #you need to use two underscores on each side to make it work! not just one underscore on each side!!
        return ("Surname: {}, Forename: {}, College: {}".format(self.surname,self.forename,self.college))
    
    def __lt__(self,other): #less than operator //remember double underscore!
        
        #N.B. python strings are immutable- you can't overwrite the value BUT you can reassign the value.
        #Objects of built-in types like (int, float, bool, str, tuple, unicode) are immutable. 
        #Objects of built-in types like (list, set, dict) are mutable. Custom classes are generally mutable.
        
        #can use <> operators on strings. Python compares the ASCII values of characters.
        if self.surname < other.surname:
            #N.B. if one is part of the other, e.g. free and freedom, then the longer one will be greater.
            return True
        elif self.surname == other.surname:
            if self.forename < other.forename:
                return True
            elif self.forename == other.forename:
                return True #doesn't really matter, completely identical names.
            else:
                return False
        else:
            return False
        
        #it appears you don't need to define __gt__ or __eq__ as long as __lt__ is completely defined.
#https://stackoverflow.com/questions/11705371/python-sort-custom-class-without-use-of-key-argument
```

<!-- grade_id: cell-ce8aafb3e627dd86, points: 0, grade: true, solution: false -->
```python
s0 = StudentEntry("Bloggs", "Andrea", 1996, 1, "Churchill", "ab1001")       
s1 = StudentEntry("Reali", "John", 1997, 1, "Corpus Christi")       
s2 = StudentEntry("Bacon", "Kevin", 1996, 1, "Newnham")
s3 = StudentEntry("Bacon", "Alexander", 1996, 1, "Queens")

assert s0 < s1
assert s0 > s2
assert s3 < s2
assert s0.age() ==  datetime.date.today().year - 1996 
assert s1.age() ==  datetime.date.today().year - 1997
assert str(s1) == "Surname: Reali, Forename: John, College: Corpus Christi"

# Test sorting
s = [s0, s1, s2, s3]
s.sort()
for earlier, later in zip(s, s[1:]):
    assert earlier.surname <= later.surname
    if earlier.surname == later.surname:
        assert earlier.forename <= later.forename
```

```python
j1 = StudentEntry("Bloggs", "Andrea", 1996, 1, "Churchill", "ab1001")       
j2 = StudentEntry("Bloggs", "Af", 1997, 1, "Corpus Christi")   
j3 = StudentEntry("Clap", "Andrea", 1996, 1, "Churchill", "ab1001")       
j4 = StudentEntry("Doll", "John", 1997, 1, "Corpus Christi")
j5 = StudentEntry("Bloggs", "Andrea", 1997, 1, "Churchill", "ab1001") 
j = [j1,j2,j3,j4,j5]
print(sorted(j))
```

Output:
```
[Surname: Bloggs, Forename: Af, College: Corpus Christi, Surname: Bloggs, Forename: Andrea, College: Churchill, Surname: Bloggs, Forename: Andrea, College: Churchill, Surname: Clap, Forename: Andrea, College: Churchill, Surname: Doll, Forename: John, College: Corpus Christi]
```

```python

```